<a href="https://colab.research.google.com/github/NSPFD/quantum-computing/blob/main/assign_classicalgates.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Sunil Karki"
COLLABORATORS = "Alone"
```

---

# Homework 1 Assignment


The places where you have enter code are marked with `# YOUR CODE HERE`.

## Question 1 (8 points)

Complete the function `OR3` that implements the 3-qubit to 1-qubit $OR$ operator. This operator acts on 4 qubits in the state $|q_0,q_1,q_2,q_3\rangle$ and maps it to $|q_0,q_1,q_2,q_3 \oplus (q_0 \text{ or } q_1 \text{ or } q_2)  \rangle$.

The function `OR3` has

- Inputs: None
- Returns: A `QuantumCircuit` object with 4 qubits and 1 classical bit that implements the $OR$ operator as described above.

Don't add any measurements to the returned circuit.

```python
!pip install qiskit
```

Output:
```
Requirement already satisfied: qiskit in /usr/local/lib/python3.12/dist-packages (2.4.1)
Requirement already satisfied: rustworkx>=0.15.0 in /usr/local/lib/python3.12/dist-packages (from qiskit) (0.17.1)
Requirement already satisfied: numpy<3,>=1.21 in /usr/local/lib/python3.12/dist-packages (from qiskit) (2.0.2)
Requirement already satisfied: scipy>=1.5 in /usr/local/lib/python3.12/dist-packages (from qiskit) (1.16.3)
Requirement already satisfied: dill>=0.3 in /usr/local/lib/python3.12/dist-packages (from qiskit) (0.3.8)
Requirement already satisfied: stevedore>=3.0.0 in /usr/local/lib/python3.12/dist-packages (from qiskit) (5.8.0)
Requirement already satisfied: typing-extensions in /usr/local/lib/python3.12/dist-packages (from qiskit) (4.15.0)
```

```python
from qiskit import QuantumCircuit

def OR3():
    # Initialize a circuit with 4 quantum qubits and 1 classical bit
    qc = QuantumCircuit(4, 1)

    # Inputs: q0, q1, q2 (indices 0, 1, 2)
    # Target: q3 (index 3)

    # YOUR CODE HERE
    # 1. Invert the input qubits so we can look for the "all zeros" condition
    qc.x([0, 1, 2])

    # 2. Invert the target qubit before checking, so we can flip it back
    # if the "all zeros" state is detected
    qc.x(3)

    # 3. Apply a Multi-Controlled X gate.
    # If q0, q1, and q2 were originally all 0, they are now all 1,
    # which triggers the MCX and flips the target (q3) back.
    qc.mcx([0, 1, 2], 3)

    # 4. Uncompute the X gates to restore the control qubits and target logic
    qc.x([0, 1, 2])

    return qc
```

<!-- grade_id: cell-517fe01d3e0bd3be, points: 0, grade: false, solution: true -->
```python
from qiskit import QuantumCircuit
from qiskit.circuit.library.standard_gates import XGate
from qiskit.primitives import StatevectorSampler

def OR3():
    circuit = QuantumCircuit(4)
    # Do not modify anything above this line

    # YOUR CODE HERE
    # 1. Invert the control qubits (0, 1, 2) to check for the "all zeros" state
    circuit.x([0, 1, 2])

    # 2. Set up the target qubit (3) negation
    circuit.x(3)

    # 3. Apply the Multi-Controlled X gate (TOFFOLI-3 variant)
    # This flips qubit 3 back if and only if qubits 0, 1, and 2 are all 1s (originally 0s)
    circuit.mcx([0, 1, 2], 3)

    # 4. Restore the control qubits back to their original input states
    circuit.x([0, 1, 2])

    # Do not modify anything below this line
    return circuit
```

In the following cell, there is some code that will help you make sure your solution is correct. You can modify this code as you wish. You may also create more cells and write more code to test your own solution.

```python
def test(input):
    circuit = QuantumCircuit(4, 1)

    # implement the input state with a suitable use of X gates.
    for i in range(3):
        if input[i]=='1':
            circuit.x(i)

    # Apply the OR3 gate you created
    circuit.compose(OR3(), range(4), inplace=True)

    circuit.measure(3, 0)

    sampler = StatevectorSampler()
    job = sampler.run([circuit], shots=100)
    result = job.result()[0]
    counts = result.data.c.get_counts()
    return counts

test('101')
```

Output:
```
{'1': 100}
```

# How autograding will work in this workshop
In the next cell, there are some tests (the `assert` statements) that will be used to *grade* if your solution is correct. If executing the next cell results in an error, your implementation is incorrect, and you will not get any marks.

There are also some hidden tests, not available to you, that will test your solution, once your submit this file back on Canvas. If any of those fail, you will not receive any marks. Otherwise, you will gain full marks on this question.

<!-- grade_id: cell-12b1e3e97b88002b, points: 2, grade: true, solution: false -->
```python
def Grade_test(input):
    circuit = QuantumCircuit(4, 1)

    for i in range(3):
        if input[i] == '1':
            circuit.x(i)

    circuit.compose(OR3(), range(4), inplace=True)
    circuit.measure(3, 0)

    sampler = StatevectorSampler()
    job = sampler.run([circuit], shots=100)
    result = job.result()[0]
    return result.data.c.get_counts()


counts = Grade_test('100')
assert counts['1']==100
counts2 = Grade_test('000')
assert counts2['0']==100
```

<!-- grade_id: cell-74f92b92e2a40b2e, points: 6, grade: true, solution: false -->
```python
# hidden tests in this cell will be used for grading.
```